<!--NOTEBOOK_HEADER-->
*This notebook contains material from [cbe-xx258](https://ndcbe.github.io/cbe-xx258);
content is available [on Github](https://github.com/ndcbe/cbe-xx258).*


<!--NAVIGATION-->
< [1.0 Python Primer](https://ndcbe.github.io/cbe-xx258/01.00-Python-Primer.html) | [Contents](toc.html) |<p><a href="https://colab.research.google.com/github/ndcbe/cbe-xx258/blob/master/docs/01.01-Jupyter-Notebooks.ipynb"> <img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a><p><a href="https://ndcbe.github.io/cbe-xx258/01.01-Jupyter-Notebooks.ipynb"> <img align="left" src="https://img.shields.io/badge/Github-Download-blue.svg" alt="Download" title="Download Notebook"></a>

# 1.1 Welcome to Jupyter Notebooks and Vocareum

## 1.1.1 Overview

During the next 10 minutes, we'll get acquainted with working in Jupyter notebooks and submitting assignments in Vocareum.

## 1.1.2 Quick Python Introduction

### 1.1.2.1 Variables

In Python, variables do not need to be declared before they are used. In other words, Python automatically decided on the correct variable type (integer, float, etc.) You can simply define a new variable using `x = 5`. This is just like MATLAB.

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
<b>Home Activity</b>: In the cell below, assign the value of 5 to the variable x. Then type Shift+Enter to execute the line.
</div>

```python

```

We can see the value of the variable using the built-in `print` function. For example, to print the value of `x`, we would use `print(x)`.

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
<b>Home Activity</b>: Print the value of x. Type Shift+Enter to execute the cell. (We'll stop reminding you to do this moving forward.)
</div>

```python

```

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
<b>Home Activity</b>: Now change the value of x to 8 and execute the cell.
</div>

```python

```

### 1.1.2.2 Execution Order

**Note**: When using Jupyter notebooks, it is very important to know that the cells can be executed out of order
(intentionally or not). The state of the environment (e.g., values of variables, imports, etc.) is defined by the execution
order.

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
<b>Home Activity</b>: To see this concept, rerun the cell above with the print statement.
</div>

You should see that the value `8` is now printed. This may seem problematic if you are used to programming in environments where the state is linked to the order of the commands as *written*, not as *executed*.

**Again, notice that the state of the environment is determined by the execution order.**
Note also that the square brackets to the left of the cell show the order that cells were executed. If you scroll to the top, you should see that the code cells show an execution order of `[1]` , `[2]` , `[5]` , and `[4]`, indicating the actual execution order.

There are some useful menu commands at the top of the Jupyter notebook to help with these problems and make sure you
retain the execution order as expected.

Some important commands to remember:
* You can clear the current state with the menu item `Kernel | Restart & Clear Output`
* It is often useful to clear the state using the menu command just described, and then execute all the lines above the currently selected cell using `Cell | Run All Above` .
* You can clear all the state and re-run the entire notebook using `Kernel | Restart & Run All`.

## 1.1.3 Vocareum Autograder

We will use the autograder in Vocareum for homework assignments, which gives you instant feedback on your work. The autograder is confirmed to allow for unlimited submissions. You just need to wait 5 minutes between submissions.

It is very important carefully read the assignment instructions. The autograder is simply executing logical checks behind the scenes. For example, if the assignment instructions say to store your final answer in the variable `final_ans` but you instead write `my_ans`, the autograder will mark it as wrong.

### Our First Autograded Submission

We want to store the **first seven elements** of the Fibonaci sequence in the list `fib`. Below is an answer with a few mistakes:

<!-- grade_id: cell-d1b0242a12e0275b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
fib = [0, 1, 1, 2, 3, 5, 8]
### END SOLUTION

fibs = [0, -1, 1, 2, 4, 5, 9, 13]
```

<!-- grade_id: sequence-length, points: 0.3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
try:
    fib
except NameError:
    print("Be sure to store the answer in `fib`.")
    raise
    
assert len(fib) == 7, "Your answer should contain SEVEN elements."
### END HIDDEN TESTS
```

<!-- grade_id: sequent-correct, points: 0.7, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
correct_answer = [0, 1, 1, 2, 3, 5, 8]

for i in range(len(correct_answer)):
    assert fib[i] == correct_answer[i], "Element "+str(i+1)+" is not correct."
### END HIDDEN TESTS
```

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
<b>Home Activity</b>: Submit this solution and look at the feedback from the autograder.
</div>

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
<b>Home Activity</b>: Fix the mistakes. You are done when all of the grading tests pass.
</div>

<!--NAVIGATION-->
< [1.0 Python Primer](https://ndcbe.github.io/cbe-xx258/01.00-Python-Primer.html) | [Contents](toc.html) |<p><a href="https://colab.research.google.com/github/ndcbe/cbe-xx258/blob/master/docs/01.01-Jupyter-Notebooks.ipynb"> <img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a><p><a href="https://ndcbe.github.io/cbe-xx258/01.01-Jupyter-Notebooks.ipynb"> <img align="left" src="https://img.shields.io/badge/Github-Download-blue.svg" alt="Download" title="Download Notebook"></a>